
## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.
Instructions

- **Download** this notebook as you would any other ipynb file 
- **Upload** to Google Colab or work locally (if you have that set-up)
- **Delete** `raise NotImplementedError()`

- **Write** your code in the `# YOUR CODE HERE` space


- **Execute** the Test cells that contain assert statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)

- **Save** your notebook when you are finished
- **Download** as a ipynb file (if working in Colab)
- **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)



# Sprint Challenge
## *Data Science Unit 4 Sprint 1*

After a week of Natural Language Processing, you've learned some cool new stuff: how to process text, how turn text into vectors, and how to model topics from documents. Apply your newly acquired skills to one of the most famous NLP datasets out there: [Yelp](https://www.yelp.com/dataset). As part of the job selection process, some of my friends have been asked to create analysis of this dataset, so I want to empower you to have a head start.  

The real dataset is massive (almost 8 gigs uncompressed). I've sampled the data for you to something more manageable for the Sprint Challenge. You can analyze the full dataset as a stretch goal or after the sprint challenge. As you work on the challenge, I suggest adding notes about your findings and things you want to analyze in the future.

## Challenge Objectives
Successfully complete all these objectives to earn full credit. 

**Successful completion is defined as passing all the unit tests in each objective.**  

Each unit test that you pass is 1 point. 

There are 5 total possible points in this sprint challenge. 


There are more details on each objective further down in the notebook.*
* <a href="#p1">Part 1</a>: Write a function to tokenize the yelp reviews
* <a href="#p2">Part 2</a>: Create a vector representation of those tokens
* <a href="#p3">Part 3</a>: Use your tokens in a classification model on yelp rating
* <a href="#p4">Part 4</a>: Estimate & Interpret a topic model of the Yelp reviews

____

# Before you submit your notebook you must first

1) Restart your notebook's Kernel

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu. 

3) Comment out the cell that generates a pyLDAvis visual in objective 4 (see instructions in that section). 
____



### Import Data

<!-- grade_id: cell-395851cd95d17235, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# Load reviews from URL
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_4/unit1_nlp/review_sample.json'

# Import data into a DataFrame named df
# YOUR CODE HERE
df = pd.read_json(data_url, lines=True)
```

<!-- grade_id: cell-cb5006475e42b8f9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert isinstance(df, pd.DataFrame), 'df is not a DataFrame. Did you import the data into df?'
assert df.shape[0] == 10000, 'DataFrame df has the wrong number of rows.'
```

## Part 1: Tokenize Function
<a id="#p1"></a>

Complete the function `tokenize`. Your function should
- accept one document at a time
- return a list of tokens

You are free to use any method you have learned this week.

```python
# Optional: Consider using spaCy in your function. The spaCy library can be imported by running this cell.
# A pre-trained model (en_core_web_sm) has been made available to you in the CodeGrade container.
# If you DON'T need use the en_core_web_sm model, you can comment it out below.
import spacy
nlp = spacy.load('en_core_web_sm')
```

```python
# Import for regex in a later function
import re
# Increasing the pandas df display for easy viewing
pd.options.display.max_colwidth = 750
df
```

Output:
```
business_id  cool                date  funny  \
0     nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1     eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2     6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3     k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4     6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   
...                      ...   ...                 ...    ...   
9995  1h3ysSuSazvXc1aeLiiOew     0 2017-10-07 10:57:15      1   
9996  Rwahe1zbFpw6VIjb5ngZeg     0 2014-01-18 15:52:52      0   
9997  8itGZAOBMiTbHKOwLuh4_Q     0 2018-08-26 02:53:21      0   
9998  A5Rkh7UymKm0_Rxm9K2PJw     0 2018-04-23 23:36:07      0   
9999  xtYiHTmunjfCN2sUaQxBjA     0 2016-06-02 23:37:51      0   

                   review_id  stars  \
0     eZs2tpEJtXPwawvHnHZIgQ      1   
1     DoQDWJsNbU0KL1O29l_Xug      4   
2     DDOdGU7zh56yQHmUnL1idQ      3   
3     LfTMUWnfGFMOfOIyJcwLVA      1   
4     zJSUdI7bJ8PNJAg4lnl_Gg      4   
..
```

<!-- grade_id: cell-3d570d5a1cd6cb64, points: 0, grade: false, solution: true -->
```python
def tokenize(doc):
# YOUR CODE HERE
    
    tokens = []
    
    doc = nlp(doc)
    
    for token in doc:
        # Filtering out punctuation, and stop words. Filtering in lemmas and case normalizion.
        if ((token.is_punct != True) and
            (token.is_lower != True) and
            (token.lemma_ != 'PRON-') and
            (token.is_stop == False)):
                tokens.append(token.lemma_)

    return tokens
```

```python
# Initializing pandarallel to speed up the tokenization process. Applying all 16 vCPUs at my disposal.
from pandarallel import pandarallel
pandarallel.initialize(progress_bar=True, nb_workers=16)
```

Output:
```
INFO: Pandarallel will run on 16 workers.
INFO: Pandarallel will use Memory file system to transfer data between the main process and workers.
```

```python
%%time
df['tokens'] = df['text'].parallel_apply(tokenize)
```

Output:
```
VBox(children=(HBox(children=(IntProgress(value=0, description='0.00%', max=625), Label(value='0 / 625'))), HB…
CPU times: user 878 ms, sys: 300 ms, total: 1.18 s
Wall time: 23.4 s
```

```python
df
```

Output:
```
business_id  cool                date  funny  \
0     nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1     eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2     6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3     k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4     6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   
...                      ...   ...                 ...    ...   
9995  1h3ysSuSazvXc1aeLiiOew     0 2017-10-07 10:57:15      1   
9996  Rwahe1zbFpw6VIjb5ngZeg     0 2014-01-18 15:52:52      0   
9997  8itGZAOBMiTbHKOwLuh4_Q     0 2018-08-26 02:53:21      0   
9998  A5Rkh7UymKm0_Rxm9K2PJw     0 2018-04-23 23:36:07      0   
9999  xtYiHTmunjfCN2sUaQxBjA     0 2016-06-02 23:37:51      0   

                   review_id  stars  \
0     eZs2tpEJtXPwawvHnHZIgQ      1   
1     DoQDWJsNbU0KL1O29l_Xug      4   
2     DDOdGU7zh56yQHmUnL1idQ      3   
3     LfTMUWnfGFMOfOIyJcwLVA      1   
4     zJSUdI7bJ8PNJAg4lnl_Gg      4   
..
```

```python
# Need to convert this tokens from a list to string for our vectorizer model later on. Will get errors if we don't.
def listintostring(string):
    '''
    This function converts lists into strings
    '''
    str = " "
    return (str.join(string))
```

```python
df['token_list'] = df['tokens'].apply(listintostring)
```

```python
df
```

Output:
```
business_id  cool                date  funny  \
0     nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1     eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2     6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3     k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4     6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   
...                      ...   ...                 ...    ...   
9995  1h3ysSuSazvXc1aeLiiOew     0 2017-10-07 10:57:15      1   
9996  Rwahe1zbFpw6VIjb5ngZeg     0 2014-01-18 15:52:52      0   
9997  8itGZAOBMiTbHKOwLuh4_Q     0 2018-08-26 02:53:21      0   
9998  A5Rkh7UymKm0_Rxm9K2PJw     0 2018-04-23 23:36:07      0   
9999  xtYiHTmunjfCN2sUaQxBjA     0 2016-06-02 23:37:51      0   

                   review_id  stars  \
0     eZs2tpEJtXPwawvHnHZIgQ      1   
1     DoQDWJsNbU0KL1O29l_Xug      4   
2     DDOdGU7zh56yQHmUnL1idQ      3   
3     LfTMUWnfGFMOfOIyJcwLVA      1   
4     zJSUdI7bJ8PNJAg4lnl_Gg      4   
..
```

<!-- grade_id: cell-02da164f6fbe730a, points: 1, grade: true, solution: false -->
```python
'''Testing'''
assert isinstance(tokenize(df.sample(n=1)["text"].iloc[0]), list), "Make sure your tokenizer function accepts a single document and returns a list of tokens!"
```

## Part 2: Vector Representation
<a id="#p2"></a>
1. Create a vector representation of the reviews (i.e. create a doc-term matrix).
2. Write a fake review and query for the 10 most similar reviews, print the text of the reviews. Do you notice any patterns?
    - Given the size of the dataset, use `NearestNeighbors` model for this. 

<!-- grade_id: cell-0e96491cb529202c, points: 0, grade: false, solution: true -->
```python
%%time
# Create a vector representation of the reviews 
# Name that doc-term matrix "dtm"

# YOUR CODE HERE
from sklearn.feature_extraction.text import TfidfVectorizer
# Tuning our vectorizer model
vect = TfidfVectorizer(
                       stop_words='english',
                       ngram_range=(1,2),    # Allowing for both unigrams AND bigrams
                       max_features=10000    # Not allowing more than 10k features/dimensions in our model
                       )
```

Output:
```
CPU times: user 1.05 s, sys: 1.13 s, total: 2.18 s
Wall time: 424 ms
```

<!-- grade_id: cell-0e96491cb529202c, points: 0, grade: false, solution: true -->
```python
# Fitting our model and assigning its output as a DTM to an appropriate name
dtm = vect.fit_transform(df['token_list'])
```

```python
# Converting our DTM to a pandas DataFrame for convenience
dtm = pd.DataFrame(dtm.toarray(), columns=vect.get_feature_names())
dtm
```

Output:
```
00  00 00  00 10  00 12  00 15  00 20  00 30  00 50  00 obviously  \
0     0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   
1     0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   
2     0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   
3     0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   
4     0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   
...   ...    ...    ...    ...    ...    ...    ...    ...           ...   
9995  0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   
9996  0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   
9997  0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   
9998  0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   
9999  0.0    0.0    0.0    0.0    0.0    0.0    0.0    0.0           0.0   

      00 pm  ...  zucchini  zucchini sticks  zumba  zumba zumba  zupas  \
0       0.0  ...    
```

<!-- grade_id: cell-3d5bc610a8ec6b24, points: 0, grade: false, solution: true -->
```python
# Create and fit a NearestNeighbors model named "nn"
from sklearn.neighbors import NearestNeighbors

# YOUR CODE HERE
nn = NearestNeighbors(n_neighbors=10, algorithm='ball_tree')  # Using ball_tree to measure distance of points
nn.fit(dtm)  # Fitting our DTM to our KNN model
```

Output:
```
NearestNeighbors(algorithm='ball_tree', n_neighbors=10)
```

```python
# sample a doc from dtm to use as our query point
doc_index = 0
# using integer location(iloc) on the document index in order to sample them
doc = [dtm.iloc[doc_index].values]

# Query Using kneighbors 
neigh_dist, neigh_index = nn.kneighbors(doc)
```

```python
neigh_dist
```

Output:
```
array([[0., 1., 1., 1., 1., 1., 1., 1., 1., 1.]])
```

```python
neigh_index
```

Output:
```
array([[  0, 154, 137, 146,  92, 109,  71,  13,  70,  37]])
```

<!-- grade_id: cell-c43704dcff67e99b, points: 1, grade: true, solution: false -->
```python
'''Testing.'''
assert nn.__module__ == 'sklearn.neighbors._unsupervised', ' nn is not a NearestNeighbors instance.'
assert nn.n_neighbors == 10, 'nn has the wrong value for n_neighbors'
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python
# Create a fake review and find the 10 most similar reviews

# YOUR CODE HERE
fake_review = ["Recently I doubled my workforce and was in need of solid, well priced vehicles in a short amount of time. The area in which I live has a limited selection of quality affordable vehicles. I exhausted my efforts purchasing vehicles locally. I was forced to search for vehicles outside of Humboldt county. Many fruitless searches finally led me to Go Best Auto sales in southern California. The vehicles looked to be in good shape according to the online photos and the prices were reasonable. Car prices were also in line with Kelly blue book, as well as other online resources (car gurus, Edmunds, etc.). I was skeptical due to many years of owning my delivery business and purchasing dozens of delivery vehicles over the past 17 years. I researched and compared similar vehicles in multiple cities and dealerships. It became very clear in a short amount of time that Shawn had the best selection, price, and professionalism available to my business. We spoke over the phone and he answered all of my questions, (which I had many), honestly and adeptly. I wasn't sure if I could afford all of the vehicles outright and he provided me with contact info for Jamie (@Cambridge capital), their in house financing. They too were patient and forthright with my needs. I explained to Shawn the need to expedite the vehicle sales (due to my rapid expansion), as well as the necessity of transport to my area in northern California. He worked with me to keep costs low as well as having my vans shipped to me in a very short period of time after the purchase of vehicles. Shawn was especially honest about the timeframe in which he could reasonably ship the vans to my business in Humboldt. He didn't want to promise something that he couldn't deliver on. It turns out that he found a reputable truck transporter in record time and had my vans delivered to me the next day (which just so happened to be a couple of days before I absolutely needed them). The Van's I bought were low miles, fleet maintained, government owned, and lacking in defects. The 2005 E350 Ford arrived with brand new tires and the 2006 Chevy Express 2500 had various aftermarket extras. Both vehicles were immaculately detailed and in perfect working order. So far so good mechanically, and I have been using them excessively during the holiday peak season rush. I will more than likely be expanding in the future and Shawn at Best Auto sales will be the person I call to purchase vehicles from. I can't say enough to Shawn's honesty, caring, attention to detail, efficiency, professionalism and dedication to providing me with quality vehicles. Thanks for being top notch and helping me with my business needs! I while heartedly reccomend Best Auto Sales for your commercial auto needs."]
```

```python
# Going through the entire process of making a new, separate dataframe with our fake review inserted in
df_new = pd.DataFrame(df['text'])
```

```python
# Inserting our fake review to the new dataframe
df_new.loc[len(df_new.index)] = fake_review
```

```python
df_new
```

Output:
```
text
0                                                                                                                                                                                                                                                BEWARE!!! FAKE, FAKE, FAKE....We also own a small business in Los Alamitos, CA and received what looked like a legitimate bill for $70 with an account number and all.  I called the phone number listed (866) 273-7934.  The wait time on hold said 20 minutes and to leave a message.  I could not get a live person on the phone no matter what number I selected.  I left a very FIRM message that I would be contacting the BBB and my attorney regarding their company trying to scam businesses. This has to be illegal!!!!!
1                                                                                                                                                                                                                                             
```

```python
dtm_new = vect.fit_transform(df_new['text'])
```

```python
dtm_new = pd.DataFrame(dtm_new.toarray(), columns=vect.get_feature_names())
dtm_new
```

Output:
```
00  00 pm  000  00am  00pm   07   10  10 00  10 10  10 12  ...  \
0      0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   
1      0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   
2      0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   
3      0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   
4      0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   
...    ...    ...  ...   ...   ...  ...  ...    ...    ...    ...  ...   
9996   0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   
9997   0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   
9998   0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   
9999   0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   
10000  0.0    0.0  0.0   0.0   0.0  0.0  0.0    0.0    0.0    0.0  ...   

       yummy food  yup  zen  zero  zero stars  zone  zoo  zucchini  était  été  
0             0.0  0.0  0.0   0.0     
```

```python
# passing in and fitting our new dataframe's DTM into our KNN model
nn.fit(dtm_new)
```

Output:
```
NearestNeighbors(algorithm='ball_tree', n_neighbors=10)
```

```python
# That's our fake review, right there!
df_new.iloc[10000]
```

Output:
```
text    Recently I doubled my workforce and was in need of solid, well priced vehicles in a short amount of time. The area in which I live has a limited selection of quality affordable vehicles. I exhausted my efforts purchasing vehicles locally. I was forced to search for vehicles outside of Humboldt county. Many fruitless searches finally led me to Go Best Auto sales in southern California. The vehicles looked to be in good shape according to the online photos and the prices were reasonable. Car prices were also in line with Kelly blue book, as well as other online resources (car gurus, Edmunds, etc.). I was skeptical due to many years of owning my delivery business and purchasing dozens of delivery vehicles over the past 17 years. I research...
Name: 10000, dtype: object
```

```python
# sampling our fake review from dtm_new to use as our query point
doc_index = 10000
# using integer location(iloc) on the document index in order to sample them
doc = [dtm_new.iloc[doc_index].values]

# Query Using kneighbors 
neigh_dist, neigh_index = nn.kneighbors(doc)
```

```python
neigh_dist
```

Output:
```
array([[0.        , 1.        , 1.        , 1.15798223, 1.21736316,
        1.27353589, 1.28115937, 1.29209967, 1.29209985, 1.29318906]])
```

```python
neigh_index
```

Output:
```
array([[10000,  6311,  6204,  5923,   487,  7759,  1588,  1645,  4951,
         7095]])
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python
# For Loop for comparing our fake review to similar reviews
for doc in neigh_index:
    print(df_new.iloc[doc])
```

Output:
```
text
10000  Recently I doubled my workforce and was in need of solid, well priced vehicles in a short amount of time. The area in which I live has a limited selection of quality affordable vehicles. I exhausted my efforts purchasing vehicles locally. I was forced to search for vehicles outside of Humboldt county. Many fruitless searches finally led me to Go Best Auto sales in southern California. The vehicles looked to be in good shape according to the online photos and the prices were reasonable. Car prices were also in line with Kelly blue book, as well as other online resources (car gurus, Edmunds, etc.). I was skeptical due to many years of owning my delivery business and purchasing dozens of delivery vehicles over the past 17 years. I research...
6311                                                                                                                                                                                                                                          
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python

```

## Part 3: Classification
<a id="#p3"></a>
Your goal in this section will be to predict `stars` from the review dataset. 

1. Create a pipeline object with a sklearn `CountVectorizer` or `TfidfVector` and any sklearn classifier.
    - Use that pipeline to train a model to predict the `stars` feature (i.e. the labels). 
    - Use that Pipeline to predict a star rating for your fake review from Part 2. 



2. Create a parameter dict including `one parameter for the vectorizer` and `one parameter for the model`. 
    - Include 2 possible values for each parameter
    - **Use `n_jobs` = 1** 
    - Due to limited computational resources on CodeGrader `DO NOT INCLUDE ADDITIONAL PARAMETERS OR VALUES PLEASE.`
    
    
3. Train the entire pipeline with a GridSearch
    - Name your GridSearch object as `gs`

```python
# Splitting our data into a feature matrix and target vector
X = df['token_list']
y = df['stars']
```

```python
# Getting our yelp review column with fake review added ready for testing
X_test = df_new['text']
```

<!-- grade_id: cell-e2beb0252d274bba, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import Pipeline
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import GridSearchCV

# Name the gridsearch instance "gs"

# YOUR CODE HERE
# Using tfidf vect and KNN for our models' pipeline
vector = TfidfVectorizer()
knnc = KNeighborsClassifier(n_jobs=-1)

pipe = Pipeline([('vect', vector),
                 ('clf', knnc)])
# Doing some statistical trimming with max_df to our vector. Applying distance and uniformity based weights to our knn
params = {
          "vect__max_df": (0.5, 0.95),
          "clf__weights": ['distance', 'uniform']
          }

gs = GridSearchCV(
                  pipe,
                  params,
                  n_jobs=-1,       # Using all our CPU resources on this
                  cv=5,            # 5 cross validation folds where fifths of our data are trained and validated on each other
                  verbose=2
                  )

gs.fit(X, y)
```

Output:
```
Fitting 5 folds for each of 4 candidates, totalling 20 fits
GridSearchCV(cv=5,
             estimator=Pipeline(steps=[('vect', TfidfVectorizer()),
                                       ('clf',
                                        KNeighborsClassifier(n_jobs=-1))]),
             n_jobs=-1,
             param_grid={'clf__weights': ['distance', 'uniform'],
                         'vect__max_df': (0.5, 0.95)},
             verbose=2)
```

```python
# Predicting with our trained, gridsearch tuned model
predict = gs.predict(fake_review)
list(predict)
```

Output:
```
[4]
```

<!-- grade_id: cell-d07134c6fe5d056e, points: 1, grade: true, solution: false -->
```python
# Visible Testing
prediction = gs.predict(["I wish dogs knew how to speak English."])[0]
assert prediction in df.stars.values, 'You gs object should be able to accept raw text within a list. Did you include a vectorizer in your pipeline?'
```

## Part 4: Topic Modeling

Let's find out what those yelp reviews are saying! :D

1. Estimate a LDA topic model of the review text
    - Set num_topics to `5`
    - Name your LDA model `lda`
2. Create 1-2 visualizations of the results
    - You can use the most important 3 words of a topic in relevant visualizations. Refer to yesterday's notebook to extract. 
3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

When you instantiate your LDA model, it should look like this: 

```python
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )

```

__*Note*__: You can pass the DataFrame column of text reviews to gensim. You do not have to use a generator.

## Note about  pyLDAvis

**pyLDAvis** is the Topic modeling package that we used in class to visualize the topics that LDA generates for us.

You are welcomed to use pyLDAvis if you'd like for your visualization. However, **you MUST comment out the code that imports the package and the cell that generates the visualization before you submit your notebook to CodeGrade.** 

Although you should leave the print out of the visualization for graders to see (i.e. comment out the cell after you run it to create the viz). 

```python
from gensim import corpora
# Due to limited computationalresources on CodeGrader, use the non-multicore version of LDA 
from gensim.models.ldamodel import LdaModel
import gensim
import re
```

Output:
```
/opt/conda/lib/python3.7/site-packages/gensim/similarities/__init__.py:15: UserWarning: The gensim.similarities.levenshtein submodule is disabled, because the optional Levenshtein package <https://pypi.org/project/python-Levenshtein/> is unavailable. Install Levenhstein (e.g. `pip install python-Levenshtein`) to suppress this warning.
  warnings.warn(msg)
```

### 1. Estimate a LDA topic model of the review tex

<!-- grade_id: cell-66331a185ff52f15, points: 0, grade: false, solution: true -->
```python
# Remember to read the LDA docs for more information on the various class attirbutes and methods available to you
# in the LDA model: https://radimrehurek.com/gensim/models/ldamodel.html

# don't change this value 
num_topics = 5

# use tokenize function you created earlier to create tokens 

# create a id2word object (hint: use corpora.Dictionary)
id2word = corpora.Dictionary(df['tokens'] )

# create a corpus object (hint: id2word.doc2bow)
corpus = [id2word.doc2bow(text) for text in df['tokens']]

# instantiate an lda model

# YOUR CODE HERE
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=59,
               num_topics=num_topics,
               passes=1
               )
```

#### Testing

<!-- grade_id: cell-5a3c181311134fa9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert lda.get_topics().shape[0] == 5, 'Did your model complete its training? Did you set num_topics to 5?'
```

#### 2. Create 1-2 visualizations of the results

<!-- grade_id: cell-9b043e992fbd218c, points: 0, grade: false, solution: true -->
```python
# import seaborn as sns
# import matplotlib.pyplot as plt
# import pyLDAvis
# import pyLDAvis.gensim_models

# # Use pyLDAvis (or a ploting tool of your choice) to visualize your results 

# # YOUR CODE HERE
# pyLDAvis.enable_notebook()
# vis = pyLDAvis.gensim_models.prepare(lda, corpus, id2word)
# vis
```

Output:
```
/opt/conda/lib/python3.7/site-packages/ipykernel/ipkernel.py:283: DeprecationWarning: `should_run_async` will not call `transform_cell` automatically in the future. Please pass the result to `transformed_cell` argument and any exception that happen during thetransform in `preprocessing_exc_tuple` in IPython 7.17 and above.
  and should_run_async(code)
PreparedData(topic_coordinates=              x         y  topics  cluster       Freq
topic                                                
2      0.193589 -0.092777       1        1  31.973465
3      0.026288  0.133618       2        1  26.078441
0     -0.042186  0.112253       3        1  19.188631
4      0.010765 -0.048528       4        1  13.507068
1     -0.188456 -0.104566       5        1   9.252395, topic_info=            Term          Freq         Total Category  logprob  loglift
0                 13398.000000  13398.000000  Default  30.0000  30.0000
19          \n\n   8017.000000   8017.000000  Default  29.0000  29.0000
56       
```

#### 3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

Due to poor cleaning, my topic model doesn't return decent topic words. I need to fix my tokenizing function to improve it. Unfortunately, I'm short on time. 

Analyzing what few good topic words my model has, its clear that the overlapping of 2 & 3 indicates that 3 is probably a subtopic of 2. 1, 4, & 5 are quite a ways away from 2 & 3 indicating that those topics probably have nothing to do with one another. Yet, when one hovers over each circle, some mention of food, a major city, and some adjective related to food.